# Prime Example

This notebook is a showcase for the Grader Service.

## First N primes

Implement the function below that takes an integer parameter (n) and returns a list of the first n prime numbers starting from 2.

<!-- grade_id: cell-a4b81a457c13e87d, points: 0, grade: false, solution: true -->
```python
def get_n_primes(n):
    ### BEGIN SOLUTION
    primes = []
    c = 2
    while n > 0:
        if all(map(lambda v: c % v != 0, primes)):
            primes.append(c)
        c += 1
        if len(primes) == n:
            return primes
    return primes
    ### END SOLUTION
```

<!-- grade_id: cell-9726d865732d68d8, points: 1, grade: true, solution: false -->
```python
assert get_n_primes(3) == [2, 3, 5]
assert get_n_primes(0) == []
```

<!-- grade_id: cell-585c88cd902849ed, points: 2, grade: true, solution: false -->
```python
# Hidden tests
### BEGIN HIDDEN TESTS
assert get_n_primes(1000)[-10:] == [7841, 7853, 7867, 7873, 7877, 7879, 7883, 7901, 7907, 7919]
### END HIDDEN TESTS
```

### Does this function also provide an efficient way to find the nth prime instead of all n primes? Why or why not?

<!-- grade_id: cell-705a6f92659c95b9, points: 1, grade: true, solution: true -->


```python

```